Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

If you find your kernel dead, please stop the process first. And then, restart your server.

Please stop your server before you log out.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jack Deng"
ID = "l630003010"
```

---

# Function Basics

*<font color="grey">Dr. Dyce Jing ZHAO<br>
Computer Science Programme,<br>
BNU-HKBU United International College</font>*

In this chapter, we are going to learn the basic syntax of *Python functions*. Functions help break our program into smaller and modular chunks. As our program grows larger and larger, functions make it more organized and manageable. It is a good habit to define and use functions. More details will be introduced in later chapters.

## Syntax of Function

The syntax of function is

```python
def function_name(parameters):
	"""docstring"""
	statement(s)
```

Below is a list of differences between a Python function and a C function:
* The keyword `def` is used to define a function in Python.
* Lines and indentation are used to indicate the function block, i.e., no bracets ({}).
* Python does not define the *return type*, and, in the body of a function, the `return` statement is optional.
* Python function has an optional documentation string (docstring) which describes what the function does.

## Function Example

A function is not executed until it is called. The following code defines and calls a function which does not return a value.

```python
# function greet(name) prints greeting message to name
def greet(name):
    """This function greets to
    the person passed in as
    parameter"""
    print("Hello, " + name + ". Nice to meet you!")
    
# function greet(name) is called
greet(input("What is your name? "))
```

And the following code defines and calls a function which returns a string.

```python
# function greet(name) returns a greeting message to name
def greet(name):
    """This function returns a
    greeting message to the person
     passed in as parameter"""
    return "Hello, " + name + ". Nice to meet you!"

# function greet(name) is called
msg = greet(input("What is your name? "))
print(msg)
```

## Docstring

The first string after the function header is called the *docstring* and is short for documentation string. It is used to explain in brief, what a function does.

Although optional, documentation is a good programming practice. In both of the two examples above, we have docstrings immediately below the function header. We generally use triple quotes so that docstring can extend up to multiple lines. This string is available to us as `__doc__` attribute of the function.

For example:

Try running the following into the Python shell to see the output.

```python
def greet(name):
    """This function returns a
    greeting message to the person
     passed in as parameter"""
    return "Hello, " + name + ". Nice to meet you!"

print(greet.__doc__)
```

## The return statement

The `return` statement is used to exit a function and go back to the place from where it was called. The syntax of return is:

```python
return [expression_list]
```

Python function can return multiple values which are seperated by comma (,), but for now let's stay with returning a single value. If there is no expression in the statement or the return statement itself is not present inside a function, then the function will return the `None` object. For example, the following code 

```python
# function didNothhing() does nothing
def didNothing():
    """This function did nothing"""
    pass
    
# when called, function didNothhing() returns None
print(didNothing())
```

## Tasks

<span class="task">Task 1</span>: In the code cell below, complete the function `absv`  which returns the absolute value of `num`.

<!-- grade_id: cell-f580fe52c6dc2e8f, points: 0, grade: false, solution: true -->
```python
# Code for Task 1
def absv(num):
    # YOUR CODE HERE
    return num if num > 0 else -num
```

<!-- grade_id: cell-b4bc4b139411d03c, points: 3, grade: true, solution: false -->
```python
assert absv(-10)== 10
assert absv(20) == 20
```

<span class="task">Task 2</span>: In the code cell below, complete the function  `gcd` which returns the greatest common divisor (GCD) of two natural numbers `x` and `y`.

<!-- grade_id: cell-6364763ce68863d9, points: 0, grade: false, solution: true -->
```python
# Code for Task 2
def gcd(x, y):
    # YOUR CODE HERE
    if x < y:
        x, y = y, x
    if y == 0:
        return x
    return y if x % y == 0 else gcd(y, x % y)
```

<!-- grade_id: cell-7ff921dfcc8b1360, points: 3, grade: true, solution: false -->
```python
assert gcd(112,230) == 2
assert gcd(11,220) == 11

```

<span class="task">Task 3</span>: In the code cell below, complete the function `checkPasswd` which checks whether the string `passwd` is a valid password or not. A password is valid if

* It contains no less than 8 characters
* It contains at least 1 number, 1 lower-case letter and 1 upper-case letter

The function returns `True` if the `passwd` is valid and `False` otherwise.

For example, `checkPasswd('aA12331222')` returns `True` while `checkPasswd('bbb1212121')` returns `False`.

<!-- grade_id: cell-5bfc51a9a469c38c, points: 0, grade: false, solution: true -->
```python
# Code for Task 3
def checkPasswd(passwd):
    # YOUR CODE HERE
    hasUpper = False
    hasLower = False
    hasNumber = False
    upperArr = [chr(i) for i in range(ord('A'), ord('Z') + 1)]
    lowerArr = [chr(i) for i in range(ord('a'), ord('z') + 1)]
    numArr = [str(i) for i in range(10)]
    if (len(passwd) >= 8):
        for char in passwd:
            if char in upperArr:
                hasUpper = True
            if char in lowerArr:
                hasLower = True
            if char in numArr:
                hasNumber = True
                
        return hasUpper and hasLower and hasNumber
    return False
```

<!-- grade_id: cell-889e463a212e78b0, points: 4, grade: true, solution: false -->
```python
assert checkPasswd('aA12331222') == True
assert checkPasswd('bbb1212121') == False
```