Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

If you find your kernel dead (e.g,, "In[*]" appears before a code cell), please interrupt the kernel (press "&#11035;").


Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", delete the statement "raise NotImplementedError()" as well as your name and collaborators below:

```python
NAME = "Jack Deng"
ID = "l630003010"
```

---

# Python List

*<font color="grey">Dr. Dyce Jing ZHAO<br>
Computer Science Programme,<br>
BNU-HKBU United International College</font>*

Python offers several datatypes which manage a sequence of data elements. List is one of the most frequently used.

## Create a list

To create a list, you may place all the list items (elements) inside a pair of bracket([]), seperated by commas (,). Follows are some notes about the declaration of a list.

* You don't have to specify the size (capacity) of the list.
* You may put items of different types in the same list.
* You may put a list as an item of another list.
* You may create an empty list.

The code cell below show some typical statements that create a list.

```python
# empty list
my_list = []

# list of integers
my_list = [1, 2, 3]

# list with mixed datatypes
my_list = [1, "Hello", 3.4]

# nested list
my_list = ["mouse", [8, 4, 6], ['a']]
```

## Access a list

### List index
We can use the index operator ([]) to access an item in a list. For a list of $n$ items, the index starts from $0$ and ends at $n-1$. The rules of list index are similar to those of C. 

Nested list is accessed using nested indexing, like the multi-dimensional array in C. Follows are some examples.

```python
my_list = ['p','r','o','b','e']
# Output: p
print(my_list[0])

# Output: o
print(my_list[2])

# Output: e
print(my_list[4])

# Error! Only integer can be used for indexing
# my_list[4.0]

# Nested List
n_list = ["Happy", [2,0,1,5]]

# Nested indexing

# Output: a
print(n_list[0][1])    

# Output: 5
print(n_list[1][3])
```

### Negative indexing

Python allows negative indexing for its sequences. The index of -1 refers to the last item, -2 to the second last item and so on. Therefore, for a list of  $n$ items, the negative index starts from $-1$ and ends at $-n$.

Try the examples below of negative index.

```python
my_list = ['p','r','o','b','e']

# Output: e
print(my_list[-1])

# Output: p
print(my_list[-5])
```

## List operation

### Slice a list

We can access a range of items in a list by using the slicing operator (:). Note that the sub-list *includes* the starting index but *excludes* the ending index. Follows is an example:

```python
my_list = ['p','y','t','h','o','n','3']
# elements 3rd to 5th
short_list = my_list[2:5]
print(short_list)

# elements beginning to 2nd
print(my_list[:-5])

# elements 6th to end
print(my_list[5:])

# elements beginning to end
print(my_list[:])
```

The index and the negative index of a list can be visualized below. Slicing a list is done by specifying the left (inclusive) and right (exclusive) boundary of a range.

|List: |p|y|t|h|o|n|3|
|---|---|---|---|---|---|---|---|
|Index:|0|1|2|3|4|5|6|
|Negative Index:|-7|-6|-5|-4|-3|-2|-1|

### Insert or update elements

List are mutable, meaning, their elements can be changed unlike string or integer. We can use the assignment operator (=) to change the value of its items. Follows is an example which shows how to change a single element and how to change a range of elements:

```python
# mistake values, should be odd numbers: 1, 3, 5, 7
odd = [2, 4, 6, 8]

# change the 1st item    
odd[0] = 1            

# Output: [1, 4, 6, 8]
print(odd)

# change 2nd to 4th items
odd[1:4] = [3, 5, 7]  

# Output: [1, 3, 5, 7]
print(odd)                   
```

We can add *one* item to the end of a list using `append()` or add several items using `extend()` method. Follows is an example:

```python
odd = [1, 3, 5]

odd.append(7)

# Output: [1, 3, 5, 7]
print(odd)

odd.extend([9, 11, 13])

# Output: [1, 3, 5, 7, 9, 11, 13]
print(odd)
```

### Delete elements

The following methods allows the deletion of list elements:

* `del`: It deletes one element at a given *index* or multiple elements at a given *index range*. This method can be used to delete the entire list. The code below demonstrates the use of `del`. Note that it throws *NameError* because the last `print` statement tries to access a list that is already deleted.

```python
my_list = ['p','r','o','b','l','e','m']

# delete one item
del my_list[2]

# Output: ['p', 'r', 'b', 'l', 'e', 'm']     
print(my_list)

# delete multiple items
del my_list[1:5]  

# Output: ['p', 'm']
print(my_list)

# delete entire list
del my_list       

# Error: List not defined
print(my_list)
```

* `remove`: It deletes one element which matches a given *value*. If multiple matches exist in the list, only the first one will be deleted.

```python
my_list = ['b', 'a', 'b', 'y']

# delete the element whose value is 'a'
my_list.remove('a')
print(my_list)

# There are two 'b' in the list and only the first one is deleted
my_list.remove('b')
print(my_list)
```

* `pop`: It removes and *returns* the one item at a given *index*. The *last item* is popped if index is not provided. <span class="hl">You may use `append` for push and `pop` for pop if you use the list as a *stack*.</span> Follows illustrates the use of pop.

```python
my_list = ['p','r','o','b','l','e','m']

# Output: 'r'
print(my_list.pop(1))

# Output: ['p', 'o', 'b', 'l', 'e', 'm']
print(my_list)

# Output: 'm'
print(my_list.pop())

# Output: ['p', 'o', 'b', 'l', 'e']
print(my_list)
```

* `clear`: It deletes all the elements in the list. Note that this method is different from `del`: it empties a list, but does not delete it entirely 
(you can still print it!). Follows is an example of clear.

```python
my_list = ['p','r','o','b','l','e','m']

my_list.clear()

# Output: []
print(my_list)
```

* Finally, we can also delete items in a list by assigning an empty list to a slice of elements. The following code piece shows the procedure.

```python
my_list = ['p','r','o','b','l','e','m']

# ['o', 'b', 'l'] are deleted from the list
my_list[2:5] = []

# Output: ['p', 'r', 'e', 'm']
print(my_list)
```

### Python list methods

Methods of the list data type are provided below. You can check <a href="https://www.programiz.com/python-programming/methods/list">programiz</a> for detail.

* `append()` - <a href="https://www.programiz.com/python-programming/methods/list/append" title="Python List append()">Add an element to the end of the list</a>
* `extend()` - <a href="https://www.programiz.com/python-programming/methods/list/extend" title="Python List extend()">Add all elements of a list to the another list</a>
* `insert()` - <a href="https://www.programiz.com/python-programming/methods/list/insert" title="Python List index()">Insert an item at the defined index</a>
* `remove()` - <a href="https://www.programiz.com/python-programming/methods/list/remove" title="Python List remove()">Removes an item from the list</a>
* `pop()` - <a href="https://www.programiz.com/python-programming/methods/list/pop" title="Python List pop()">Removes and returns an element at the given index</a>
* `clear()` - <a href="https://www.programiz.com/python-programming/methods/list/clear" title="Python List clear()">Removes all items from the list</a>
* `index()` - <a href="https://www.programiz.com/python-programming/methods/list/index" title="Python List index()">Returns the index of the first matched item</a>
* `count()` - <a href="https://www.programiz.com/python-programming/methods/list/count" title="Python List count()">Returns the count of number of items passed as an argument</a>
* `sort()` - <a href="https://www.programiz.com/python-programming/methods/list/sort" title="Python List sort()">Sort items in a list in ascending order</a>
* `reverse()` - <a href="https://www.programiz.com/python-programming/methods/list/reverse" title="Python List reverse()">Reverse the order of items in the list</a>
* `copy()` - <a href="https://www.programiz.com/python-programming/methods/list/copy" title="Python List copy()">Returns a shallow copy of the list</a>

## List comprehension: Elegant way to create a list

List comprehension is an elegant and concise way to create new list from an existing list in Python.

List comprehension consists of an expression followed by for statement inside square brackets.

Here is an example to make a list with each item being increasing power of 2.

```python
pow2 = [2 ** x for x in range(10)]

# Output: [1, 2, 4, 8, 16, 32, 64, 128, 256, 512]
print(pow2)
```

## Other list operations

### List Membership Test

We can test if an item exists in a list or not, using the keyword `in`. Follows is an example:

```python
my_list = ['p','r','o','b','l','e','m']

# Output: True
print('p' in my_list)

# Output: False
print('a' in my_list)

# Output: True
print('c' not in my_list)
```

### Iterating Through a List

Using the `in` keyword and a `for` loop we can iterate though each item in a list.

```python
for fruit in ['apple','banana','mango']:
    print("I like",fruit)
```

### Built-in Functions with List

The following functions are available for all the *iterable* data types (e.g., list, tuple, ...). Therefore, they are also available for lists.

* <a href="https://www.programiz.com/python-programming/methods/built-in/all" title="Python all()">all()</a> - Return True if all elements of the list are true (or if the list is empty).
* <a href="https://www.programiz.com/python-programming/methods/built-in/any" title="Python any()">any()</a> - Return True if any element of the list is true. If the list is empty, return False.
* <a href="https://www.programiz.com/python-programming/methods/built-in/enumerate" title="Python enumerate()">enumerate()</a> - Return an enumerate object. It contains the index and value of all the items of list as a tuple.
* <a href="https://www.programiz.com/python-programming/methods/built-in/len" title="Python len()">len()</a> - Return the length (the number of items) in the list.
* <a href="https://www.programiz.com/python-programming/methods/built-in/list" title="Python list()">list()</a> - Convert an iterable (tuple, string, set, dictionary) to a list.
* <a href="https://www.programiz.com/python-programming/methods/built-in/max" title="Python max()">max()</a> - Return the largest item in the list.
* <a href="https://www.programiz.com/python-programming/methods/built-in/min" title="Python min()">min()</a> - Return the smallest item in the list
* <a href="https://www.programiz.com/python-programming/methods/built-in/sorted" title="Python sorted()">sorted()</a> - Return a new sorted list (does not sort the list itself).
* <a href="https://www.programiz.com/python-programming/methods/built-in/sum" title="Python sum()">sum()</a> - Return the sum of all elements in the list.

## Tuple

Tuples can be considered as *immutable lists*. The difference between the two is that we cannot change the elements of a tuple once it is assigned whereas in a list, elements can be changed.

### Advantages of Tuple over List
Since tuples are quite similiar to lists, both of them are used in similar situations as well.

However, there are certain advantages of implementing a tuple over a list. Below listed are some of the main advantages:

* We generally use tuple for heterogeneous (different) datatypes and list for homogeneous (similar) datatypes.
* Since tuple are immutable, iterating through tuple is faster than with list. So there is a slight performance boost.
* Tuples that contain immutable elements can be used as key for a dictionary. With list, this is not possible.
* If you have data that doesn't change, implementing it as tuple will guarantee that it remains write-protected.

### Create a tuple

A tuple is created by placing all the items (elements) inside a parentheses (()), separated by comma(,). The parentheses are optional but is a good practice to write it.

A tuple can have any number of items and they may be of different types (integer, float, list, string etc.).

Follows are some examples:


```python
# empty tuple
# Output: ()
my_tuple = ()
print(my_tuple)

# tuple having integers
# Output: (1, 2, 3)
my_tuple = (1, 2, 3)
print(my_tuple)

# tuple with mixed datatypes
# Output: (1, "Hello", 3.4)
my_tuple = (1, "Hello", 3.4)
print(my_tuple)

# nested tuple
# Output: ("mouse", [8, 4, 6], (1, 2, 3))
my_tuple = ("mouse", [8, 4, 6], (1, 2, 3))
print(my_tuple)

# tuple can be created without parentheses
# also called tuple packing
# Output: 3, 4.6, "dog"

my_tuple = 3, 4.6, "dog"
print(my_tuple)

# tuple unpacking is also possible
# Below: a = 3, b = 4.6, c = "dog"
a, b, c = my_tuple
print(a)
print(b)
print(c)
```

Creating a tuple with one element is a bit tricky.

Having one element within parentheses is not enough. We will need a trailing comma to indicate that it is in fact a tuple.

```python
# only parentheses is not enough
# Output: <class 'str'>
my_tuple = ("hello")
print(type(my_tuple))

# need a comma at the end
# Output: <class 'tuple'>
my_tuple = ("hello",)  
print(type(my_tuple))

# parentheses is optional
# Output: <class 'tuple'>
my_tuple = "hello",
print(type(my_tuple))
```

### More on tuples

Read the [online tutorial](https://www.programiz.com/python-programming/tuple) to learn more about tuples.

## Tasks

<span class="task">Task 1</span>: In the code cell below, complete function `hello` which says hello to every name in `name_list`. 

For example, if *name_list = ['Mary', 'Tracy']*, the function should print:
```
Hello, Mary!
Hello, Tracy!
```

*Note*: Check [here](https://www.programiz.com/python-programming/input-output-import) for formatted `print`.

<!-- grade_id: cell-38d243e8b2870e35, points: 1, grade: true, solution: true -->
```python
# Code for Task 1
def hello(name_list):
    # YOUR CODE HERE
    for name in name_list:
        print("Hello, " + name + "!")
```

<span class="task">Task 2:</span> In the code cell below, complete function `absList` which creates and returns a list whose numbers are the absolute values of those in `num_list`.

For example, *absList([-1, 2, -4])* should return *[1, 2, 4]*.

<!-- grade_id: cell-532a0971bb1c9f91, points: 0, grade: false, solution: true -->
```python
# Code for Task 2
def absList(num_list):
    # YOUR CODE HERE
    return list(map(abs, num_list))
```

<!-- grade_id: cell-308600c2917e4701, points: 1, grade: true, solution: false -->
```python
a=[-1, 2, -4]
assert absList(a)==[1,2,4]

```

<span class="task">Task 3</span>: In the code cell below, complete the function `onlyPositive` which deletes all the zero or negative values in `num_list` and returns the updated `num_list`.

For example, if *num_list = [-2, -3, 5, -4]*, *onlyPositive(num_list)* should delete *-2*, *-3* and *-4* from `num_list` and return *[5]*.

<!-- grade_id: cell-ee64a8d0f90653e1, points: 0, grade: false, solution: true -->
```python
# Code for Task 3
def onlyPositive(num_list):
    # YOUR CODE HERE
    delList = []#Record number to delete
    for number in num_list:
        if number <= 0:
            delList.append(number)
    
    for number in delList:
        num_list.remove(number)
        
    return num_list
```

<!-- grade_id: cell-63f8b48629fd8af2, points: 2, grade: true, solution: false -->
```python
num_list = [-2, 3, 5, -4,-9,-10]
assert onlyPositive(num_list)==[3, 5]



```

<span class="task">Task 4</span>: In the code cell below, complete the function `addMatrix` which adds two matrics, `A` and `B`, and returns the sum. Note that the size of `A` and `B` is equal but is not fixed.

For example, *addMatrix([[1, 2], [3, 4]], [[1, 1], [2, 2]])* should return *[[2, 3], [5, 6]]*.

*Hint*: Use append() to add an element to the end of a list.

<!-- grade_id: cell-99f82f6cf7edaeac, points: 0, grade: false, solution: true -->
```python
# Code for Task 4
def addMatrix(A, B):
    # YOUR CODE HERE
    newMatrix = []#The result of sum of A and B
    height = len(A)#Length of first dimension of list
    width = len(A[0])#Length of second dimension of list
    for i in range(height):
        newList = []#List of result of A and B
        for j in range(width):
            newList.append(A[i][j] + B[i][j])
        newMatrix.append(newList)
    return newMatrix
```

<!-- grade_id: cell-5e829d522e0607aa, points: 2, grade: true, solution: false -->
```python
X=[[1, 2], [3, 4]]
Y=[[1, 1], [2, 2]]
assert addMatrix(X, Y) ==  [[2, 3], [5, 6]]



```

<span class="task">Task 5</span>: In the code cell below, complete the function `multiplyMatrix` which multiplies two matrics, `A` and `B`, and returns the product matrix. Note that the size of `A` and `B` is equal but is not fixed.

For example, *multiplyMatrix([[1, 2], [3, 4]], [[2, 1], [1, 2]])* should return *[[4, 5], [10, 11]]*.

*Hint*: Use append() to add an element to the end of a list.

<!-- grade_id: cell-191790e006422cf6, points: 0, grade: false, solution: true -->
```python
# Code for Task 5
def multiplyMatrix(A, B):
    # YOUR CODE HERE
    result = []#List to store all the results
    size = len(A)#Width and height of matrix
    for i in range(size):
        line = []#List to store results of a row in matrix
        for j in range(size):
            sum = 0#The result of current element
            for k in range(size):
                sum += A[i][k] * B[k][j]
            line.append(sum)
        result.append(line)
    return result
```

Output:
```
[[23, 20, 11], [47, 40, 23], [71, 60, 35]]
```

<!-- grade_id: cell-f8db74b3c9f87196, points: 2, grade: true, solution: false -->
```python
A=[[1, 2, 3], [3, 4, 5], [5, 6, 7]]
B=[[2, 1, 1], [9, 8, 5], [1, 1, 0]]
assert multiplyMatrix(A, B) == [[23, 20, 11], [47, 40, 23], [71, 60, 35]]

```

<span class="task">Task 6</span>: In the code cell below, complete the function `uniqueList` which returns a new list consisting of the unique characters in `char_list`. If repeated character exists in `char_list`, the new list only keeps its first occurrence.

For example, if *char_list = ['m', 'o', 'r', 'n', 'i', 'n', 'g', 'c', 'a', 'l', 'l']*, then *uniqueList(char_list)* should return *['m', 'o', 'r', 'n', 'i', 'g', 'c', 'a', 'l']*.

<!-- grade_id: cell-d8a2699dfc960efe, points: 0, grade: false, solution: true -->
```python
# Code for Task 6
def uniqueList(char_list):
    # YOUR CODE HERE
    index = []#List to store position to pop
    #Find the position to pop
    for i in range(1, len(char_list)):
        if char_list[i] in char_list[:i]:
            index.append(i)
      
    for i in range(len(index)):
        char_list.pop(index[i] - i)
    return char_list
```

<!-- grade_id: cell-add03ad760306fbb, points: 2, grade: true, solution: false -->
```python
char_list = ['m', 'o', 'r', 'n', 'i', 'n', 'g', 'c', 'a', 'l', 'l']
assert uniqueList(char_list) == ['m', 'o', 'r', 'n', 'i', 'g', 'c', 'a', 'l']

```