Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

If you find your kernel dead (e.g,, "In[*]" appears before a code cell), please interrupt the kernel (press "&#11035;").


Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", delete the statement "raise NotImplementedError()" as well as your name and collaborators below:

```python
NAME = "Jack Deng"
ID = "l630003010"
```

---

# Python Set

*<font color="grey">Dr. Dyce Jing ZHAO<br>
Computer Science Programme,<br>
BNU-HKBU United International College</font>*

A set is an unordered collection of items. Every element is unique (no duplicates) and must be *immutable* (which cannot be changed). However, the set itself is mutable. We can add or remove items from it.

Sets can be used to perform mathematical set operations like union, intersection, symmetric difference etc.

## Create a Set

A set is created by placing all the items inside a pair of curly braces ({}), separated by comma (,). It can have items of different types. However, a set item must be immutable (e.g., integer, float, tuple, string, ...).

Follows are valid examples that create sets:

```python
# an empty set
my_set = {}
print(my_set)

# a set of integers
my_set = {1, 2, 7}
print(my_set)

# a set of mixed data types
# (1, 2.3) is a tuple which is immutable
my_set = {1, 2.2, "hello", (1, 2.3)}
print(my_set)
```

And follows are improper examples of set creation:

```python
# set do not have duplicates, so repeated items will be ignored
# Output: {1, 2, 3, 4}
my_set = {1,2,3,4,3,2}
print(my_set)

# set cannot have mutable items
# here [3, 4] is a list (mutable)
# this will cause an error.
# TypeError: unhashable type: 'list'
my_set = {1, 2, [3, 4]}
```

In addition, you may create a set from a list using the built-in function `set()`. Follows shows the syntax.

```python
# a tuple is defined by ()
# a list is defined by []
# a set is defined by {}
# the following set is created from a list
my_list = [1, 2, 3]
my_set = set(my_list)
print(my_set)

my_str = "bravo"
my_set = set(my_str)
print(my_set)
```

Creating an empty set is a bit tricky.

Empty curly braces ({}) will make an empty *dictionary*, which will be introduced later. To make an empty set, we use the `set()` function without any argument. Follows is an example:

```python
# initialize something with {}
something = {}

# check data type of something
# Output: <class 'dict'>
print(type(something))

# initialize something with set()
something = set()

# check data type of something
# Output: <class 'set'>
print(type(something))
```

## Add Items to a Set

Sets are mutable, which means you can insert or delete an item from a set. However, set items are unordered and indexing have no meaning in a set.

We cannot access or change an element of set using indexing or slicing. Set does not support it. We can 

* add single element using the `add()` method
* and add multiple elements using the `update()` method. 

The `update()` method can take tuples, lists, strings or other sets as its argument. In all cases, duplicates are ignored.

Follows shows the use of `add()` and `update()`:

```python
# initialize my_set
my_set = {3, 1}
print(my_set)

# if you uncomment line: my_set[0]
# you will get an error
# TypeError: 'set' object does not support indexing

#my_set[0]

# add an element
# Output: {1, 2, 3}
my_set.add(2)
print(my_set)

# add multiple elements from list
# duplicates are ignored
# Output: {1, 2, 3, 4, 5}
my_set.update([2,3,4,5])
print(my_set)

# add multiple elements from tuple and set
# duplicates are always ignored
# Output: {1, 2, 3, 4, 5, 6, 7, 8}
my_set.update((7,5), {1,6,8})
print(my_set)
```

## Remove Items from a Set

A particular item can be removed from set using methods, `discard()` or `remove()`.

The only difference between the two is that, while using `discard()` if the item does not exist in the set, it remains unchanged. But` remove()` will raise an error in such condition.

The following example will illustrate this:

```python
# initialize my_set
my_set = {1, 3, 4, 5, 6}
print(my_set)

# discard an element
# Output: {1, 3, 5, 6}
my_set.discard(4)
print(my_set)

# remove an element
# Output: {1, 3, 5}
my_set.remove(6)
print(my_set)

# discard an element
# not present in my_set
# Output: {1, 3, 5}
my_set.discard(2)
print(my_set)

# remove an element
# not present in my_set
# If you uncomment the last line,
# you will get an error.
# Output: KeyError: 2

#my_set.remove(2)
```

Set supports the `pop()` method. However, as set items are unordered, there is no way of determining which item will be popped. It is completely arbitrary.

We can also remove all items from a set using clear().

Follows is an example:

```python
# initialize my_set
# Output: set of unique elements
my_set = set("HelloWorld")
print(my_set)

# pop an element
# Output: random element
print(my_set.pop())

# pop another element
# Output: random element
my_set.pop()
print(my_set)

# clear my_set
#Output: set()
my_set.clear()
print(my_set)
```

## Set Operations

Sets can be used to carry out mathematical set operations like union, intersection, difference and symmetric difference. We can do this with operators or methods.

Let us consider the following two sets for the following operations.

```python
>>> A = {1, 2, 3, 4, 5}
>>> B = {4, 5, 6, 7, 8}
```

### Set Union
![Set Union](img/ch7.union.jpg)

As illustrated by the image above, union of `A` and `B` is a set of all elements from both sets.

Union of `A` and `B` is performed using
* `A|B`
* or `A.union(B)`

Note that this operation produces the union set but does not change set `A` or `B`. Follows is an example:

```python
# initialize A and B
A = {1, 2, 3, 4, 5}
B = {4, 5, 6, 7, 8}

# union A and B
# Output: {1, 2, 3, 4, 5, 6, 7, 8}
print(A|B)
print(A.union(B))
```

### Set Intersection

![Set Intersection](img/ch7.intersection.jpg)

As illustrated by the image above, intersection of A and B is a set of elements that are common in both sets.

Intersection of `A` and `B` is performed using
* `A&B`
* or `A.intersection(B)`

### Set Difference

![Set Difference](img/ch7.difference.jpg)

Difference of `A` and `B` (`A-B`) is a set of elements that are only in `A` but not in `B`. Similarly, `B-A` is a set of element in `B` but not in `A`.

Difference of `A` and `B` is performed using

* `A-B`
* or `A.difference(B)`

### Set Symmetric Difference

![Set Symmetric Difference](img/ch7.sdifference.jpg)

Symmetric difference of `A` and `B` is a set of elements in both `A` and `B` except those that are common in both.

Symmetric difference of `A` and `B` is performed using

* `A^B`
* or `A.symmetric_difference(B)`

Read the code below and try to predict the output for each `print()` statement.

```python
# initialize A and B
A = {1, 2, 3, 4, 5}
B = {4, 5, 6, 7, 8}

print(A&B)
print(A-B)
print(B-A)
print(A^B)
```

### Python Set Methods

There are many set methods, some of which we have already used above. Here is a list of all the methods that are available with set objects.

* <a href="https://www.programiz.com/python-programming/methods/set/add" title="Python set add()">add()</a> - Add an element to a set
* <a href="https://www.programiz.com/python-programming/methods/set/clear" title="Python set clear()">clear()</a> - Remove all elements form a set
* <a href="https://www.programiz.com/python-programming/methods/set/copy" title="Python set copy">copy()</a> - Return a shallow copy of a set
* <a href="https://www.programiz.com/python-programming/methods/set/difference" title="Python set difference()">difference()</a> - Return the difference of two or more sets as a new set
* <a href="https://www.programiz.com/python-programming/methods/set/difference_update" title="Python set difference_update">difference_update()</a> - Remove all elements of another set from this set
* <a href="https://www.programiz.com/python-programming/methods/set/discard" title="Python set discard()">discard()</a> - Remove an element from set if it is a member. (Do nothing if the element is not in set)
* <a href="https://www.programiz.com/python-programming/methods/set/intersection" title="Python set intersection()">intersection()</a> - Return the intersection of two sets as a new set
* <a href="https://www.programiz.com/python-programming/methods/set/intersection_update" title="Python set intersection_update()">intersection_update()</a> - Update the set with the intersection of itself and another
* <a href="https://www.programiz.com/python-programming/methods/set/isdisjoint" title="Python set isdisjoint()">isdisjoint()</a> - Return `True` if two sets have a null intersection
* <a href="https://www.programiz.com/python-programming/methods/set/issubset" title="Python set issubset()">issubset()</a> - Return `True` if another set contains this set
* <a href="https://www.programiz.com/python-programming/methods/set/issuperset" title="Python set issuperset()">issuperset()</a> - Return `True` if this set contains another set
* <a href="https://www.programiz.com/python-programming/methods/set/pop" title="Python set pop()">pop()</a> - Remove and return an arbitary set element. Raise `KeyError` if the set is empty
* <a href="https://www.programiz.com/python-programming/methods/set/remove" title="Python set remove()">remove()</a> - Remove an element from a set. If the element is not a member, raise a `KeyError`
* <a href="https://www.programiz.com/python-programming/methods/set/symmetric_difference" title="Python set symmetric_difference()">symmetric_difference()</a> - Return the symmetric difference of two sets as a new set
* <a href="https://www.programiz.com/python-programming/methods/set/symmetric_difference_update" title="Python set symmetric_difference_update()">symmetric_difference_update()</a> - Update a set with the symmetric difference of itself and another
* <a href="https://www.programiz.com/python-programming/methods/set/union" title="Python set union()">union()</a> - Return the union of sets in a new set
* <a href="https://www.programiz.com/python-programming/methods/set/update" title="Python set update()">update()</a> - Update a set with the union of itself and others

## Other Set Operations
### Set Membership Test

Set membership can be tested in a similar way List membership is tested. The keyword is `in`. Further, `for` loops can be used to iterate items in a set, too. Below is an example, note that the order of set items is arbitrary.

```python
for fruit in {'apple','banana','mango'}:
    print("I like",fruit)
```

### Built-in Functions with Set

They are THE SAME to those for a List!


* <a href="https://www.programiz.com/python-programming/methods/built-in/all" title="Python all()">all()</a> - Return True if all elements of the list are true (or if the list is empty).
* <a href="https://www.programiz.com/python-programming/methods/built-in/any" title="Python any()">any()</a> - Return True if any element of the list is true. If the list is empty, return False.
* <a href="https://www.programiz.com/python-programming/methods/built-in/enumerate" title="Python enumerate()">enumerate()</a> - Return an enumerate object. It contains the index and value of all the items of list as a tuple.
* <a href="https://www.programiz.com/python-programming/methods/built-in/len" title="Python len()">len()</a> - Return the length (the number of items) in the list.
* <a href="https://www.programiz.com/python-programming/methods/built-in/list" title="Python list()">list()</a> - Convert an iterable (tuple, string, set, dictionary) to a list.
* <a href="https://www.programiz.com/python-programming/methods/built-in/max" title="Python max()">max()</a> - Return the largest item in the list.
* <a href="https://www.programiz.com/python-programming/methods/built-in/min" title="Python min()">min()</a> - Return the smallest item in the list
* <a href="https://www.programiz.com/python-programming/methods/built-in/sorted" title="Python sorted()">sorted()</a> - Return a new sorted list (does not sort the list itself).
* <a href="https://www.programiz.com/python-programming/methods/built-in/sum" title="Python sum()">sum()</a> - Return the sum of all elements in the list.

## Python Frozenset

Frozenset is a new class that has the characteristics of a set, but its elements *cannot be changed* once assigned. <span class="hl">While tuples are immutable lists, frozensets are immutable sets.</span>

Frozensets can be created using the function `frozenset()`.

This datatype supports methods like `copy()`, `difference()`, `intersection()`, `isdisjoint()`, `issubset()`, `issuperset()`, `symmetric_difference()` and `union()`. Being immutable, it does not have method that add or remove elements.

The code below shows how a frozen set is created. A forzen set can be the member of a set, since it is immutable.

```python
iset = frozenset([1, 2, 3, 4])
mset = {1, iset, "bye"}
print(mset)
```

## Tasks

<span class="task">Task 1</span>: In the code cell below, complete the function `equalSets` which decided if two sets, `A` and `B`, are equal.

For example, *equalSets({1,2}, {2, 1})* returns `True` while *equalSets({1,2}, {2, 3})* returns `False`.

<!-- grade_id: cell-810d1a17b1f73a22, points: 0, grade: false, solution: true -->
```python
# Code for Task 1

def equalSets(A, B):
    # YOUR CODE HERE
    return len(A - B) == 0
```

<!-- grade_id: cell-2f8a7f368c5330ed, points: 2, grade: true, solution: false -->
```python
A={1,2}
B={2,1}
assert equalSets(A, B) == True
```

<span class="task">Task 2</span>: In the code cell below, complete the function `uniqueWords` which returns a set containing all the distinct words that appear in the string `sentence`. `sentence` contains (possibly repeated) words separated by whitespaces.

For example, *uniqueWords("good good study day day up")* returns *{"good", "study", "day", "up"}*.

*Note*:
* It is okay if your set is printed in a different order compared to the sample, because order does not matter in a set!
* Check the [split() method](https://www.programiz.com/python-programming/methods/string/split) of a `string` data type which divide a sentence into words.

<!-- grade_id: cell-9e6f5189017e7ec6, points: 0, grade: false, solution: true -->
```python
# Code for Task 2

def uniqueWords(sentence):
    # YOUR CODE HERE
    return set(sentence.split(" "))
```

<!-- grade_id: cell-83442e8000ebef8d, points: 3, grade: true, solution: false -->
```python
assert uniqueWords("good good study day day up") == {'day', 'good', 'study', 'up'}

```

<span class="task">Task 3:</span>: In the code cell below, complete the function `heavyOverlap` which accepts two sets, `A` and `B`, as parameters. The function returns `True` if at least half of the items in `A` appear in `B`, and `False` otherwise.

For example, if *A = {1, 2, 3}* and *B = {2, 3, 4, 5, 6, 7, 8}*, *heavyOverlap(A, B)* returns `True` because 2 out of the 3 items in `A` appear in `B`, which is more than half. However, *heavyOverlap(B, A)* returns `False`.

<!-- grade_id: cell-59a8ee35800cf20b, points: 0, grade: false, solution: true -->
```python
# Code for Task 3

def heavyOverlap(A, B):
    # YOUR CODE HERE
    return round(len(A) / 2) <= len(A & B) 
```

Output:
```
True
```

<!-- grade_id: cell-69b359897b49ad76, points: 2, grade: true, solution: false -->
```python
A = {1, 2, 3}
B = {2, 3, 4, 5, 6, 7, 8}
assert heavyOverlap(A, B) == True


```

<span class="task">Task 4:</span>: In the code cell below, complete the function `objectAttribute` which accepts a parameter, `obj`, of immutable data type. The function returns a set containing three piece of attributes of `obj`:

* the *id*
* the data type
* the value

For example, *objectAttribute(5)* may return *{1967550048, 5, int}* while *objectAttribute((1, 32))* may return *{(1, 32), 1939399570056, tuple}*.

*Note*: The id depends on the running of the program.

<!-- grade_id: cell-585e7a6d492c6e84, points: 3, grade: true, solution: true -->
```python
# Code for Task 4

def objectAttribute(obj):
    # YOUR CODE HERE
    attr = set({})
    attr.add(obj)
    attr.add(id(obj))
    attr.add(type(obj))
    return attr
```

```python
objectAttribute((10, 302))

```

## Self Test

Take the quiz on [Programiz](https://www.programiz.com/python-programming/quiz/native-data-types/takehttps://www.programiz.com/python-programming/quiz/decision-making/take/) on the topic of lists, tuples and sets. You may skip the last two questions.
