Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

If you find your kernel dead (e.g,, "In[*]" appears before a code cell), please interrupt the kernel (press "&#11035;").


Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE" and delete the statement "raise NotImplementedError()" 

Don't forget write down your name and collaborators below:

```python
NAME = "Jack Deng"
ID = "l630003010"
```

---

# Python Dictionary

**If we want to store information about students, you may choose to use lists. Every list is used for one set of information such as the following.**
- ```IDs = ['001', '002', '003']```
- ```names = ['Lily', 'Mary', 'Martin']```
- ```grades = ['A', 'B', 'C']```

If you want to update or retrieve student information, how would you do it? You may use the following chunk of code.

```python
IDs = ['001', '002', '003']
names = ['Lily', 'Mary', 'Martin']
grades = ['A', 'B', 'C']

def getNameGrade(ID, nameList, gradesList):
    i = IDs.index(ID)
    name = nameList[i]
    grade = gradesList[i]
    return (name, grade)

for id in range(3):
    print(getNameGrade(IDs[id], names, grades))

```

### Disadvantages using lists for this type of problems:
- Very messy if there is different information
- So many lists to maintain
- Must always use integer index
- Must change multiple lists


### Here is a better way to solve such problem - Dictionary.

## Dictonary




### The picture below shows the difference between the List and a Dictionary.
![](Gigi_img/dict_img1.png)

** Example: a dictionary named gradesDict containing students name and grades info**

|    |   |
|--- |---|
| Lily | A |
| Mary | B |
| Martin | C |
| Tom    | B+ |


** Define the gradesDict:**
```Python 
gradesDict = {'Lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
```

### A few ways to define a dictionary:

** 1.Define an empty dictionary**
```python
# It will create an empty dictionary instead of a set, 
# although they both use curly braces "{}" to include the data.
empty_dict = {}
# or
empty_dict = dict()
```
** 2. key : value**
```python
gradesDict = {'Lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
```
** 3. convert from a list composed of tuples**
```python
gradesDict = dict([('Lily', 'A'), ('Mary', 'B'), ('Martin', 'C'), ('Tom', 'B+')]) 
```
** 4. dict(key = value)**
```python
gradesDict = dict(Lily = 'A', Mary = 'B', Martin = 'C', Tom= 'B+')
```
**Note:This way cannot work if the type of the key is integer. For example, "a = dict(1 = 'one', 2 = 'two')" will cause an error. **

** 5. using zip()**
```python
gradesDict = dict(zip(['Lily', 'Mary', 'Martin', 'Tom'], ['A', 'B', 'C', 'B+']))

```




** Task 1: ** Add three more students and their grades into the gradesDict above, and then print out the dictionary.

|    |   |
|--- |---|
| Mathew | D |
| Lynne | C+ |
| Andrew | C |

<!-- grade_id: cell-7ca78110ca9580e7, points: 0, grade: false, solution: true -->
```python
gradesDict = {'Lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
# gradesDict = dict([('Lily','A'),('Mary','B'),('Martin','C'),('Tom','B+')]) 
# gradesDict = dict(Lily = 'A', Mary = 'B', Martin = 'C', Tom= 'B+')
# gradesDict = dict(zip(['Lily','Mary','Martin','Tom'],['A','B','C','B+']))
# you code for Task 1 here.
# YOUR CODE HERE
gradesDict['Andrew'] = 'C'
gradesDict['Mathew'] = 'D'
gradesDict['Lynne'] = 'C+'
print(gradesDict)
```

<!-- grade_id: cell-22772f4142460e8a, points: 3, grade: true, solution: false -->
```python
assert gradesDict['Mathew'] == 'D'
assert gradesDict['Lynne'] == 'C+'

assert gradesDict['Andrew'] == 'C'
```

### How to look up the dictionary?
- Looks up the key
- Returns the value associated with the key
- Get an error if the key isn't found in the dictionary

** Task 2: ** Run the following code and observe the output. Then change the dictionary to gradesDict, and print out Martin's grade.

<!-- grade_id: cell-2c3d01ca93457b66, points: 1, grade: true, solution: true -->
```python
## code for Task 2
dict1 = {'Name':'Mary', 'Age':19}
print("dict1['Name']:", dict1['Name'])
print("dict1['Age]:", dict1['Age'])

# print the record of Matin in gradesDict

# YOUR CODE HERE
gradesDict = {'Lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
print("gradesDict['Martin']:", gradesDict['Martin'])
```

Output:
```
dict1['Name']: Mary
dict1['Age]: 19
gradesDict['Martin']: C
```

### Dictionary Operations

Let's use the gradesDict again:
```python
gradesDict = {'Lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
```
- If we want to add a new student, we can use:

```python
gradesDict['Bob'] = 'A-'
```

- If we want to test if a key is in the dictionary, we write:
```python
'Lily' in gradesDict
```

** Task 3:** Run the following code and see what it will return. Write a print statement to test a name that is not in the dictionary, and observe what it will return. 

```python
# code for Task 3
gradesDict = {'Lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
print ('Lily' in gradesDict)
```

Output:
```
True
```

** To delete an entry in the dictionary, you can use:**
```python
del(gradesDict['Lily'])
```

### Comparison of  Dictionary Keys and Values

| keys | Values |
| ---  | ------ |
| must be <b>unique</b> | can contain <b>duplicates</b>|
| <b>immutable</b> type (int, float, string, tuple, bool) | and type (<b>immutable and mutable</b>) |
| Try to advoid using float type key | can be lists, even other dictionaries |

** Task 4: ** Run the following code and observe the output.

```python
# code for Task 4
gradesDict = {'lily':'A', 'Mary':'B', 'Martin':'C', 'Tom':'B+'}
print(gradesDict.keys())
print(gradesDict.values())
```

Output:
```
dict_keys(['lily', 'Mary', 'Martin', 'Tom'])
dict_values(['A', 'B', 'C', 'B+'])
```

### Comparison of List and Dictionary
| List | Dictionary |
| ---- | ------- |
| Ordered Sequence of elements | Matching Key and Values |
| Index must be an integer | Key can be any immutable type |
| Indices has an order | No order is guaranteed |
| Search element by index | Search value by the key |

** Task 5: ** Create a dictionary to hold some English words and corresponding Chinese words. Let the user enter a sentence that uses the words in the dictionary.  You program should translate that sentence into Chinese as output.

| keys | values |
| -- | -- |
| I | 我 |
| you | 你 |
| love | 爱 |
| China | 中国 |
| Zhuhai | 珠海 |
| and | 和 |
| me | 我 |

<!-- grade_id: cell-b3824cfa2a8eec42, points: 2, grade: true, solution: true -->
```python
# code for Task 5

# YOUR CODE HERE
dictionary = dict(zip(['I', 'you', 'love', 'China', 'Zhuhai', 'and', 'me']
                     , ["我", "你", "爱", "中国", "珠海", "和", "我"]))
sentence = input("Enter a sentence with words in the table: ")
for word in sentence.split(" "):
    print(dictionary[word], end="")
```


## Several built-in functions of Dictionary
### 2.1. fromkeys()

- fromkeys() is used to return a new dictionary.
- The syntax:

    - ```dict.fromkeys(seq[, value]))```

- The first parameter is required, and it should be a sequence(list type, tuple type or str type), every element in the sequence will be one key of the new dictionary; the second one is the value for each key. The second parameter is optional.

```python
dict1 = {}

dict2 = dict1.fromkeys((1, 2, 3))  # dict2=dict1.fromkeys(1,2,3) will return an error
print("dict1 is: ", dict1) # Output: {}
print("dict2 is: ", dict2) # Output:{1: None, 2:None, 3:None}， dict1 won't be changed

dict3 = dict1.fromkeys((1, 2, 3), ("Integer"))
print("dict3 is: ", dict3)  # Output:{1:'Integer',2:'Integer',3:'Integer'}

dict4 = dict1.fromkeys((1, 2, 3), ("one", "two", "three"))
print("dict4 is: ", dict4)  # Output: {1: ('one', 'two', 'three'), 2: ('one', 'two', 'three'), 3: ('one', 'two', 'three')}
```

### 2.2. keys(), values(), items()
- keys() returns **all the keys** in the dictionary
- values() returns **all the values** in the dictionary
- items() returns **all the items(composed of "key:value")** in the dictionary

```python
my_dict = dict()
my_dict = my_dict.fromkeys(range(5), 'UIC')
print(my_dict)
print("All the keys are: ", my_dict.keys())
print("All the values are: ", my_dict.values())
print("All the items are: ", my_dict.items())
```

### 2.3 copy()

``` copy() ``` is used to copy a dictionary.

The following example shows the difference between ```b = a.copy()``` and ```b = a```(a and b are both dictionary type).

```python
# dictC and dictD are the same dictionary, dictE is a new one.
dictC = dict(zip(['I', 'II', 'III'], [1, 2, 3]))
dictD = dictC
dictE = dictC.copy()

print("dictC is: ", dictC, ", and dictC's address is: ", id(dictC))
print("dictD is: ", dictD, ", and dictD's address is: ", id(dictD))
print("dictE is: ", dictE, ", and dictE's address is: ", id(dictE))
```

### 2.4 get(), setdefault()
** To obtain the value of one item in the dictionary, there are three methods. Let's make dict_a a dictionary.**


- ```dict_a.get(key)``` ---will return the value of key; if key doesn't exist, it will return "None", while the original dictionary remains unchanged 
- ```dict_a.setdefault(key)```---will return the value of key; if key doesn't exist, it will return "None", and a new item "key2:None" will be added into the orginal dictionary
- ```dict_a[key]```--- will return the value of key; if key doesn't exist, it will cause a keyerror, the orginal dictionary stays unchanged.

So, ```get()``` and ```setdefault()``` look much more friendly.

**The two examples below show the differences.**

```python

dict_a = dict(zip(['telephone', 'computer', 'desk', 'book', 'pen'],['t', 'c', 'd', 'b', 'p']))
x = dict_a['computer']
print(x, " is short for 'computer'")
print(dict_a) 


# when the corresponding key is out of the range in the given dictionary,
# it will return a keyerror, uncomment the following lines to test
#y= my_dict['ereaser']
#print(y)
```

```python

mydict = dict(zip(['I','II','III','IV','V'],[1,2,3,4,5]))

a_copy = mydict.copy() # copy a new dictionary
b = mydict.setdefault('IV')
c = mydict.get('IV')

print("The original dictionary is: ", mydict)
print("Using setdefault(), we get the fouth value is: ", b)
print("Using get(), we get the fouth value is: ", c)

d = mydict.get('VI')
print("\n\nThe orginal dictionary now becomes: ", mydict)
print("The sixth value is:", d)  # Output: The sixth value is: None

e = mydict.setdefault('VI')  
print("\n\nThe orginal dictionary now becomes: ", mydict)
print("The sixth value is:%s, and item 'VI: None' is added" % e)  # Output: The sixth value is: None, item "'VI': None" is added
```

### 2.5 clear()

The ```clear()``` function is used to clear a dictionary. Some students may think it is also okay to use ```dict1 = {} ``` to clear an existed dictionary. Are they the same?

The following example will tell you the answer.

```python
# This example try to use "dictA={}" to clear a dictionary

dictA = dict(I = 1, II = 2, III = 3)
dictB = dictA
print("dictA is: ", dictA, "\ndictB is: ", dictB)
print("The address of dictA is: ", id(dictA))
print("The address of dictB is: ", id(dictB))

dictA = {}
print("\n\ndictA is: ", dictA, "\ndictB is: ", dictB)

```

As shown in the example above, dictA and dictB both point to the same dictionary: ```{'I': 1, 'II': 2, 'III': 3}```. When we use ```dictA = {}``` to make dictA= point to an empty dictionary, whreas the original dictionary still exists in the memory. The dictionary cannot be truely cleaned out. Instead, clear() can do that.

The following example will help us to understand that.

```python
# This example try to use "dictA.clear()" to clear a dictionary

dictA = dict(I = 1, II = 2, III = 3)
dictB = dictA
print("dictA is: ", dictA, "\ndictB is: ", dictB)
print("The address of dictA is: ", id(dictA))
print("The address of dictB is: ", id(dictB))

dictA.clear()
print("\n\ndictA is: ", dictA,"\ndictB is: ", dictB)
```

### 2.6 pop(), popitem()

- The syntax(a is a dictionary, for example):
    - ```a.pop(key)```
    - ```a.popitem()```
- The difference between them is: ```a.pop(key)``` will return the item with the appointed key; however, as elements in a dictionary are disordered, we cannot use index to locate one item, so ```a.popitem()``` will return an item <b>randomly</b>.
- See the following example:

```python
a = {'banana':'b', 'apple':'a', 'grape':'g', 'pumpkin':'p'}
copy_a = a.copy()
print("a is: ", a)

b = a.pop('grape')
print("b is: ", b, "\nthe type of the b is: ", type(b))

# c = a.pop('grape')
# If we want to pop out item that has been popped out, for example, c = a.pop('grape'), 
# it will return an keyerror

print("after the first pop, a is: ", a)

print("\n\n", copy_a.popitem())
print(copy_a.popitem())
print(copy_a.popitem())
```

### 2.7 update() 

**Two ways to change the items in a dictionary**(dict_a, for example):

- ```dict_a.update(key = "something")``` 
- ```dict_a[key] = "something else"``` 

See the following example:

```python
dict_a = {'telephone': 't', 'computer': 'c', 'desk': 'd', 'books': 'b', 'pens': 'p'}

dict_a.update(pens = 'pe') # update the value of an existing item
print(dict_a)

dict_a.update(cellphone = 'cp') # update/add the value of a new item
print(dict_a)


print("\n\n")
dict_a['desk'] = 'ds' # use dict_a[key]="something else" to update the value of an existing item
print(dict_a)

dict_a['earphone'] = 'ep' # use dict_a[key]="something else" to update/add the value of a new item
print(dict_a)
```

### 2.8 To learn more Built In Functions...
By using:

```dir(dict)```

we can know more built-in functions of dictionary, learn more about that by google or baidu. 

```python
dir(dict)
```

**Task 6**

Cici, the TA of the workshop II course, asks you for help to develop a simple Score Management System. It should include the following functions:

- **1. Use ```upd_rec(stu_name, stu_score)``` to update or add a record. if the student name dosen't exist, just add the record into the dictionary, ```score_dbs```; if it exists, it should ask the user "Change the record?(Y or N)", if the answer is "Y" or "y", then update the corresponding record; if the answer is "N" or "n", then remain the dictionary unchanged. See the following picture.**
![](Gigi_img/dict_img3.png)


<!-- grade_id: cell-ec1ce642a219f417, points: 3, grade: true, solution: true -->
```python
# score_dbs, a global dict type, is also the 'database' used to save students' information
score_dbs = dict()

def upd_rec(stu_name, stu_score):
    global score_dbs
    # YOUR CODE HERE
    if stu_name not in score_dbs:#Check if the name is in score_dbs
        score_dbs[stu_name] = stu_score
        return
    ch = input("Change the record?(Y or N)")#Record the input of users
    if ch == 'Y' or ch == 'y':
        score_dbs[stu_name] = stu_score
    
# you can add your test here.
upd_rec("Terra",91) 
upd_rec("Ben",64) 
upd_rec("Cannon",72) 
# above are three new
# next, raise an changed example
upd_rec("Terra",85)
# raise an unchanged example
upd_rec("Ben",51)
print(score_dbs)
```

Output:
```
Change the record?(Y or N)y
Change the record?(Y or N)n
{'Terra': 85, 'Ben': 64, 'Cannon': 72}
```

- **2. Inserting records one by one is really troublesome, sometimes Cici wants to insert several records at one time. Please program to realize that. **(Tips: call ```upd_rec(stu_name, stu_score)``` by loop)
![](Gigi_img/dict_img2.png)

<!-- grade_id: cell-0ed70acaeafcacf5, points: 2, grade: true, solution: true -->
```python
stuName = ["Mathew", "Cannon", "Sue", "Leslie", "Daniel", "Calvin"]
stuScore = [65, 80, 77, 50, 95, 46]
score_dbs = dict()

def upd_rec(stu_name, stu_score):
    global score_dbs
    # YOUR CODE HERE
    if stu_name not in score_dbs:#Check if the name is in score_dbs
        score_dbs[stu_name] = stu_score
        return
    ch = input("Change the record?(Y or N)")#Record the input of users
    if ch == 'Y' or ch == 'y':
        score_dbs[stu_name] = stu_score
# YOUR CODE HERE
for (name, score) in zip(stuName, stuScore): #Get students' information in loop
    upd_rec(name, score)
    
print(score_dbs)
```

Output:
```
{'Mathew': 65, 'Cannon': 80, 'Sue': 77, 'Leslie': 50, 'Daniel': 95, 'Calvin': 46}
```

- **3: Use ```del_rec(stu_name)``` to delete a record. if the student name dosen't exist, then return a tip "The record doesn't exist." and remain the dictionary unchanged; if it exist, it should ask the user "Delete the record?(Y or N)", if the answer is "Y" or "y", then delete the corresponding record; if the answer is "N" or "y", then remain the dictionary unchanged. See the following picture.**
![](Gigi_img/dict_img4.png)

<!-- grade_id: cell-65a8fc3d7cdf7659, points: 3, grade: true, solution: true -->
```python
score_dbs = {'Mathew': 65, 'Sue': 77, 'Cannon': 80, 'Leslie': 50, "Daniel": 95, 'Calvin': 46}
def del_rec(stu_name):

    # YOUR CODE HERE
    if (stu_name not in score_dbs):#Check if student name in score_dbs
        print("The record doesn't exist.")
        return
    select = input("Delete the record?(Y or N)")
    if select == 'Y' or select == 'y':#Check user's input
        score_dbs.pop(stu_name)
    
    
# you can add your test here.
del_rec('Cici') # raise an unexisted example
del_rec('Mathew') # raise an existed example but I don't want to delete it
del_rec('Sue')   # raise an existed example and I want to delete it
print(score_dbs)
```

Output:
```
The record doesn't exist.
Delete the record?(Y or N)n
Delete the record?(Y or N)y
{'Mathew': 65, 'Cannon': 80, 'Leslie': 50, 'Daniel': 95, 'Calvin': 46}
```

- **4: Use ```lkp_rec(stu_name)``` to lookup a record. if the student name dosen't exist, then return a tip "The record doesn't exist." and remain the dictionary unchanged; if it exist, then return the corresponding record. **

<!-- grade_id: cell-822b91c0d08a8d31, points: 0, grade: false, solution: true -->
```python
def lkp_rec(stu_name):
    # YOUR CODE HERE
    if stu_name not in score_dbs:#Find if the name is in score_dbs
        print("The record doesn't exist.")
        return (stu_name, None)
    return (stu_name, score_dbs[stu_name])
```

<!-- grade_id: cell-1b7fa7bd24f03084, points: 2, grade: true, solution: false -->
```python
score_dbs = dict()
upd_rec("Terra", 91) 
upd_rec("Ben", 64) 
upd_rec("Cannon", 73) 

assert lkp_rec('Ben') == ('Ben', 64)
assert lkp_rec('Cici') == ('Cici', None)


```

- **5: Use ```lkp_all_recs()``` to return all the records in the dictionary.**

<!-- grade_id: cell-894f8586aca3eedd, points: 0, grade: false, solution: true -->
```python
def lkp_all_recs():

    # YOUR CODE HERE
    return score_dbs

# you can add your test here.
lkp_all_recs()
```

Output:
```
{'Ben': 64, 'Cannon': 73, 'Terra': 91}
```

<!-- grade_id: cell-173c3adaf584437f, points: 1, grade: true, solution: false -->
```python
score_dbs = dict()
upd_rec("Terra", 91) 
upd_rec("Ben", 64) 
upd_rec("Cannon", 73) 

assert lkp_all_recs() == {'Ben': 64, 'Cannon': 73, 'Terra': 91}

```

- **6: Use ```aver_score()``` to return the average score of all the students; if the dictionary is empty, then return 0.**

<!-- grade_id: cell-174590a8d049780e, points: 0, grade: false, solution: true -->
```python

def aver_score():

    # YOUR CODE HERE
    result = 0
    count = 0#Count number of students
    if len(score_dbs) == 0:
        return 0
    for key in score_dbs:
        result += score_dbs[key]
        count += 1
    return result / count
    
# you can add your test here
```

<!-- grade_id: cell-dfc6334cc55fe37c, points: 2, grade: true, solution: false -->
```python
score_dbs = dict()
upd_rec("Terra", 91) 
upd_rec("Ben", 64) 
upd_rec("Cannon", 73) 
avg = aver_score()
assert avg == 76

```

<!-- grade_id: cell-77c5e105429b2f62, points: 1, grade: true, solution: false -->
```python
score_dbs = dict()
avg = aver_score()
assert avg == 0
```